# Onlineaktivität 1: Sortieren und Suchen
(Gruppenaufgabe)

```julia
String[][] phoneBook = {
    /* firstname, surname, phone number */
    {"Courtney","Acosta","(01) 5565 4523"},
    {"Yardley","Lucas","(03) 4677 0681"},
    {"Jennifer","Webb","(06) 8604 6467"},
    {"Leroy","Schmidt","(06) 8109 6804"},
    {"Aimee","Barry","(07) 1700 6664"},
    {"Charlotte","Chambers","(05) 8261 0680"},
    {"Kameko","Guthrie","(07) 8222 2843"},
    {"Branden","Gillespie","(01) 2779 7205"},
    {"Karyn","Morin","(01) 0658 1271"},
    {"Joel","Carter","(01) 5860 7796"},
    {"Jamalia","Ferguson","(08) 5290 9722"},
    {"Alice","Watts","(04) 4748 4864"},
    {"Phillip","Nieves","(09) 8511 0785"},
    {"Wyatt","Perkins","(05) 7904 0076"},
    {"Mona","Walter","(06) 3666 7682"},
    {"Aquila","Hanson","(09) 2137 1735"},
    {"Ursa","Roach","(03) 1816 3766"},
    {"Priscilla","Hall","(03) 9265 7056"},
    {"Ruth","Delacruz","(05) 6939 7409"},
    {"Nora","Hanson","(02) 4855 4456"},
    {"Lamar","Serrano","(03) 0738 6158"},
    {"Dustin","Frank","(08) 7699 7903"},
    {"Stacy","Henry","(08) 3849 9603"},
    {"Imani","Drake","(02) 1880 8612"},
    {"Rae","Hammond","(03) 8385 1875"},
    {"Jerome","Cooke","(07) 2872 5563"},
    {"Kelsey","Livingston","(02) 7904 6218"},
    {"Magee","Collier","(02) 9145 0959"},
    {"Kasper","Bird","(05) 6394 1974"},
    {"Cruz","Compton","(07) 4523 7390"},
    {"Harper","Glover","(03) 1754 8818"},
    {"Xavier","Willis","(02) 6710 6797"},
    {"Baker","Roth","(06) 0718 1462"},
    {"Melvin","Hunter","(08) 6606 7222"},
    {"Blossom","Buck","(01) 3714 8019"},
    {"Yardley","Holland","(08) 1845 8652"},
    {"Dennis","Williams","(03) 9386 1976"},
    {"Brooke","West","(02) 8751 0567"},
    {"Timothy","Osborn","(04) 2403 3375"},
    {"Frances","Arnold","(09) 3442 6062"},
    {"Quincy","Myers","(07) 3826 9298"},
    {"Lois","Matthews","(05) 7673 8203"},
    {"Allen","Hoover","(02) 7087 6822"},
    {"Ahmed","Lott","(03) 2345 5508"},
    {"Kane","Abbott","(02) 8807 9695"},
    {"Erich","West","(03) 4634 6576"},
    {"Colin","Fulton","(03) 9714 4864"},
    {"Fuller","Cole","(04) 3758 5247"},
    {"Joy","Bridges","(09) 0487 1071"},
    {"Ivor","Fry","(05) 3127 2880"},
    {"Beau","Taylor","(06) 2715 1790"},
    {"Cooper","Head","(07) 9082 4874"},
    {"Raymond","Moreno","(09) 6468 4582"},
    {"Tana","Andrews","(01) 9283 3438"},
    {"Brian","House","(07) 0849 5313"},
    {"Clarke","Sharp","(08) 3472 7278"},
    {"Frances","Blevins","(03) 6695 8353"},
    {"Chaney","Puckett","(06) 8185 5179"},
    {"Venus","Nolan","(05) 3677 8169"},
    {"Amethyst","Cash","(07) 1721 1128"},
    {"Cadman","Torres","(05) 7036 4304"},
    {"Emerson","Rhodes","(01) 5280 2483"},
    {"Imani","Walker","(03) 0547 0139"},
    {"Forrest","Pittman","(03) 1108 0957"},
    {"Paul","Ramos","(03) 3663 2658"},
    {"Neil","Oneil","(04) 1042 1380"},
    {"Iliana","Duncan","(08) 3999 0482"},
    {"Rina","Terrell","(05) 6499 1164"},
    {"Basia","Hays","(06) 5162 9657"},
    {"Rigel","Haynes","(01) 2542 4565"},
    {"Reuben","Cochran","(03) 4138 2365"},
    {"Hayes","Sims","(01) 3214 3452"},
    {"Avye","Lopez","(05) 6148 2472"},
    {"Craig","Mcclure","(01) 7396 6409"},
    {"Benjamin","Wells","(03) 2073 2131"},
    {"Latifah","Conway","(06) 8628 1375"},
    {"Scott","Jefferson","(01) 4512 0466"},
    {"Hedwig","Torres","(07) 2765 7164"},
    {"Frances","Mcpherson","(07) 2047 9068"},
    {"Deirdre","Valencia","(01) 5635 9044"},
    {"Ima","Hopkins","(06) 0980 9352"},
    {"Lawrence","Giles","(06) 7605 2238"},
    {"Colin","Raymond","(08) 4716 9997"},
    {"Molly","Burke","(01) 3198 8000"},
    {"Desiree","Workman","(04) 9219 6651"},
    {"Yoshio","Browning","(06) 9264 8633"},
    {"McKenzie","Giles","(08) 5280 4980"},
    {"Honorato","Blanchard","(09) 9703 5253"},
    {"Clarke","Ortiz","(03) 4991 1294"},
    {"Renee","Hull","(03) 3170 0677"},
    {"Brady","Byers","(06) 7572 7911"},
    {"Judah","Washington","(06) 0614 8493"},
    {"Nora","Todd","(06) 7876 8155"},
    {"Chase","Pollard","(08) 0699 2966"},
    {"Cynthia","Puckett","(02) 2065 5235"},
    {"Quincy","Gould","(03) 7455 2115"},
    {"Aristotle","Vaughan","(04) 3533 1912"},
    {"Wynter","Hamilton","(02) 1056 7540"},
    {"Barclay","Pena","(09) 9589 5648"},
    {"Armando","Pearson","(06) 1665 7102"}
};
```

Output:
```
[6] invokelatest at ./essentials.jl:709 [inlined]
 [7] eventloop(::ZMQ.Socket) at /opt/julia/packages/IJulia/rWZ9e/src/eventloop.jl:8
 [8] (::IJulia.var"#15#18")() at ./task.jl:356
```

## Aufgabe 1

Gegegeben ist eine unsortierte Datenstruktur (ein zweidimensionales String-Feld: `String[][]`), welche ein Telefonbuch repräsentiert (sh. `phoneBook`). Jeder Telefonbucheintrag enthält Nachname, Vorname und Telefonnummer. Weiters gegeben ist eine Sortierungsfunktion `sort`. Das Ziel ist die aufsteigende Sortierung der Telefonbucheinträge nach Nachname. Bei gleichen Nachnamen entscheidet aufsteigend der Vorname.

1. Diskutieren Sie zunächst die folgenden drei Fragen:
  * Die Implementierung welches Sortieralgorithmus liegt vor?
  * Erfolgt die Sortierung `in-place` oder `out-place`?
  * Welche Defekte sind vorhanden? 
  * Wie sind die Defekte in der Implementierung zu beheben?
2. Beheben Sie die gefundenen Defekte in der Implementierung von `sort`.
3. Diskutieren Sie die folgenden weiterführenden Fragen (auf Basis der korrekten Implementierung):
  * Warum wird `String.compareTo()` verwendet, und nicht nur ein einfacher Ausdruck mit Vergleichsoperator, z.B.: `>`, `<`?
  * Warum wird in der inneren `for`-Schleife die Zählvariable `j` um `-i` korrigiert?
  * Gibt es eine Möglichkeit, ohne die Variable `temp` auszukommen?

```julia
/* Original, defective implementation */
String[][] sort(String[][] pb) {
    for (int i = 0; i < pb.length; i++) {
        for (int j = 0; j < pb.length - i - 1; j++) { 
            if ((pb[j][1] + pb[j][0]).compareTo(pb[j + 1][1] + pb[j + 1][0]) < 0) {
                String[] temp = pb[j + 1];
                pb[j] = pb[j];
                pb[j + 1] = temp;
            }
        }
    }
    return pb;
}
```

<!-- grade_id: cell-82785b1fd912a53c, points: 0, grade: false, solution: true -->
```julia
/* Your corrected implementation */

// YOUR CODE HERE
String[][] sort(String[][] pb) {
    for (int i = 0; i < pb.length; i++) {
        for (int j = 0; j < pb.length - i - 1; j++) {         
            if ((pb[j][1] + pb[j][0]).compareTo(pb[j + 1][1] + pb[j + 1][0]) > 0) {
                String[] temp = pb[j];
                pb[j] = pb[j + 1];
                pb[j + 1] = temp;
            }
        }
    }
    
    return pb;
}

var result = sort(phoneBook);
```

<!-- grade_id: cell-226f50d1b2828c1a, points: 4, grade: true, solution: false -->
```julia
/* DO NOT DELETE OR ALTER THIS CELL! */

assert result.length == phoneBook.length : "T1.A1";

```

## <span style="color:red"> Antworten: </span>
 
1. 
  + Es liegt ein `Bubble Sort`Algorithmus vor.
  + Durch die Benutzung von dem String "temp", erfolgt die sortierung `"out-place"`
  + Defekte: "<" anstatt ">" in der if Schleife, if Schleife wurde ausgeführt wenn nicht mehr sortiert werden musste
              `pb[j] = pb[j + 1]`, der alphabetisch niedrigere Wert wird jetzt an die richtige Position gesetzt.
 
 
2. 
  * Lösung im obigen Code
 
3. 
  * Erklärung `.compareTo()` : 
   ".compareTo" sortiert Strings nach dem Alphabet; "<" oder ">" ist nur auf Zahlen anwendbar <br> 
   Bsp: `int result = StringA.compareTo(StringB)` <br>
   `result` ist: <br> `kleiner als 0 wenn` : String B alphabethisch größer als String A <br>`0 wenn` : String A und B den selben alphabethischen Value haben<br> `größer als 0` : wenn String A einen größeren alphabethischen Value als String B hat
  * `j-i` in der `for-Schleife` : `j` wird um `i` korrigiert, da die höchsten Werte dann bereits sortiert sind 
  * `temp` in der `if-Schleife` : Beim konkreten Bubble-Sort `Nein` , bei anderen Algorithmen wie zB. "Selection-Sort" findet die Vertauschung `in-place`statt - hier wäre es möglich ohne `temp`auszukommen 

## Task 2

Ausgehend von einem sortierten Telefonbuch (`result`), implementieren Sie die drei folgenden Such- bzw. Filterfunktionen.

*Hinweis: Sie müssen entscheiden, welche Datenstrukturen Sie innerhalb der Funktionen verwenden, um die Auswahl zu halten. Sie haben dabei alle Optionen: eine statische Datenstruktur wie etwa einen zweiten zwei-dimensionalen Array oder dynamische Datenstrukturen, die Java bietet. Achten Sie auch darauf, dass Sie die Anzahl der benötigten Iterationen und Vergleichsoperationen minimieren!*

### 2.1 `findByAreaCode`

Die Funktion `findByAreaCode` erhält als Argumente das sortierte Telefonbuch und eine Vorwahl ("area code", als String!) als Suchbegriff. Die Funktion retourniert die Auswahl des sortierten Telefonbuchs, welche Telefonnummern mit der entsprechenden Vorwahl aufweisen. Wenn keine passenden Einträge gefunden werden, wird ein leeres zweidimensionales Feld (`String[][]` mit `length` 0) zurückgeliefert.

<!-- grade_id: cell-4bee33ac595f3b5f, points: 0, grade: false, solution: true -->
```julia
String[][] findByAreaCode(String[][] pb, String areaCode) {
   // YOUR CODE HERE
    
    ArrayList<String[]> areaCodeResult = new ArrayList<String[]>();
    String searchCode = "("+areaCode+")";
    
    for (int i = 0; i < pb.length; i++) {
        if (pb[i][2].contains(searchCode)) {
            areaCodeResult.add(pb[i]);
        }
    }

    return areaCodeResult.toArray(new String[areaCodeResult.size()][3]);
}
```

<!-- grade_id: cell-9983be0ca657a536, points: 1, grade: true, solution: false -->
```julia
/* DO NOT DELETE OR ALTER THIS CELL! */

assert findByAreaCode(result, "09").length <= result.length : "T2.1.A1";


```

### 2.2 `filterByRepeatedFirstName`

Die Funktion `filterByRepeatedFirstName` erhält als Argument das sortierte Telefonbuch. Sie liefert die Auswahl jener Telefonbucheinträge, die einen Vornamen aufweisen, der im Telefonbuch (über der Sammlung aller Einträge) mehr als nur einmal vorkommt. Werden keine derartigen Einträge gefunden, wird ein leeres zweidimensionales Feld (`String[][]` mit `length` 0) zurückgeliefert.

<!-- grade_id: cell-b686b3dbaac2b3cf, points: 0, grade: false, solution: true -->
```julia
String[][] filterByRepeatedFirstName(String[][] pb) {
    // YOUR CODE HERE
    
    ArrayList<String[]> repeatedFirstNames = new ArrayList<String[]>();
    
    for (int i = 0; i < pb.length; i++) {
        String firstName = pb[i][0];
        
        for (int j = 0; j < pb.length; j++) {         
            if (pb[j][0].equals(firstName) && j != i) {
                repeatedFirstNames.add(pb[i]);
                break;
            }
        }
    }
    
    return repeatedFirstNames.toArray(new String[repeatedFirstNames.size()][3]);
}
```

<!-- grade_id: cell-43bfa4b6a687f2e7, points: 2, grade: true, solution: false -->
```julia
/* DO NOT DELETE OR ALTER THIS CELL! */

assert filterByRepeatedFirstName(result).length <= result.length : "T2.2.A1";

```

### 2.3 `filterByUniqueSurname`

Die Funktion `filterByUniqueSurname` erhält als Argument das sortierte Telefonbuch. Sie liefert die Auswahl jener Telefonbucheinträge, die einen Nachnamen aufweisen, der im Telefonbuch (über der Sammlung aller Einträge) nur einmal vorkommt. Werden keine derartigen Einträge gefunden, wird ein leeres zweidimensionales Feld (`String[][]` mit `length` 0) zurückgeliefert.

<!-- grade_id: cell-609625cef2ed0d4d, points: 0, grade: false, solution: true -->
```julia
String[][] filterByUniqueSurname(String[][] pb) {
    // YOUR CODE HERE
    
    ArrayList<String[]> uniqueSurnames = new ArrayList<String[]>();
    
    for (int i = 0; i < pb.length; i++) {
        String surName = pb[i][1];
        Boolean repeated = false;
        
        for (int j = 0; j < pb.length; j++) {         
            if (pb[j][1].equals(surName) && j != i) {
                repeated = true;
                break;
            }
        }
        
        if (repeated == false) {
            uniqueSurnames.add(pb[i]);
        }
    }
    
    return uniqueSurnames.toArray(new String[uniqueSurnames.size()][3]);
}
```

<!-- grade_id: cell-cf8894abb41ed120, points: 3, grade: true, solution: false -->
```julia
/* DO NOT DELETE OR ALTER THIS CELL! */

assert filterByUniqueSurname(result).length <= phoneBook.length : "T2.3.A1";
```