```java
import java.util.stream.*; 
```

## Assignment 4 (Group)

Für diese Aufgabe müssen Sie zuerst die Datei `students.txt` [1] aus MyLearn in Ihrer Jupyter-Umgebung ablegen; diese muss im gleichen Verzeichnis wie dieses Notebook untergebracht werden. Der folgende Codeblock liest die Daten in eine dynamische Datenstruktur ein:


[1] Sh. https://matrikel.adbk.de/matrikel

```java
BufferedReader reader = new BufferedReader(new FileReader("students.txt"));
String line = reader.readLine();
String[] arr;
int counter = 0;

List<String[]> data = new ArrayList<String[]>();

while (line != null) {
    data.add(line.split(";"));
    line = reader.readLine();
}
reader.close();
```

Mit den folgenden Zeilen 
wird eine Klasse `Student` bereitgestellt. Ein `Student` wird durch einen Namen, das Immatrikulationsjahr, eine fortlaufende Immatrikulationszahl und das Studienfeld beschrieben.

```java
class Student {
    int registrationYear;
    String serial; 
    String name;
    String field;
    Student(String rY, String s, String n, String f) {
        this.registrationYear = Integer.parseInt(rY);
        this.serial = s;
        this.name = n;
        this.field = f;
    }

    @Override
    public String toString() {
        StringBuffer sb = new StringBuffer();
        sb.append(this.registrationYear);
        sb.append("|");
        sb.append(this.serial);
        sb.append("|");
        sb.append(this.name);
        sb.append("|");
        sb.append(this.field);
        return sb.toString();
    }
}
```

## Aufgabe 1

Implementieren Sie die Methode `id`, welche einen eindeutigen Bezeichner für Student*innen darstellt. Dieser Bezeichner soll sich aus dem Einstiegsjahr und der Matrikelnummer zusammensetzen.

Vervollständigen Sie die Funktion `id` auf Basis eines `Student`. Der return-Wert soll die Vereinigung aus `year` + `serial` sein, beschränkt auf positive Integer.

<!-- grade_id: cell-4785d59176b7b418, points: 0, grade: false, solution: true -->
```java
int id(Student s) {
    String year = String.valueOf(s.registrationYear);
    String yearSerial = year + s.serial;
    
    return  Integer.valueOf(yearSerial);
}

//https://www.javatpoint.com/java-string-to-int
//https://www.freecodecamp.org/news/java-string-to-int-how-to-convert-a-string-to-an-integer/
```

<!-- grade_id: cell-e1d9e01c07d3b090, points: 1, grade: true, solution: false -->
```java

assert id(new Student("1932", "00011", "Max Mustermann", "Malerei")) == 193200011 : "A.1.1";
```

## Aufgabe 2a

Vervollständigen Sie die Funktion `N()`. Die Aufgabe von `N()` ist die Ermittlung der Größe einer Hashtabelle auf Basis der folgenden "Daumenregeln": 

- Basis der Berechnung ist die zu erwartende maximale Zahl an Schlüsseln in der Hashtabelle: `expected`;
- Die Größe soll um einen Faktor zwischen 1,3 (Minimum) und 1,4 (Maximum) größer sein als diese Basis.
- Die Größe soll eine Primzahl sein. Nutzen Sie `isPrime()` für die Implementierung von `N()`.

```java
boolean isPrime(int n) {
    int i;
    for(i=2; i<=Math.sqrt(n); i++){
        if(n % i == 0){
           return false;
        }
    }
    return true;
}
```

<!-- grade_id: cell-5ecca090b90a7b92, points: 0, grade: false, solution: true -->
```java
int N(int expected) {
    double exp = Double.valueOf(expected);
    int minimum = (int)(Math.ceil(exp * 1.3));
    int maximum = (int)(Math.floor(exp * 1.4));
    for (int i = minimum; i <= maximum; i++) {
        if (isPrime(i)) {
            return i;
        }
    } 
    return 0;
}

//https://www.delftstack.com/de/howto/java/java-round-up/
//https://www.tutorialspoint.com/java/lang/math_ceil.htm
```

<!-- grade_id: cell-e55613da15b18d1b, points: 2, grade: true, solution: false -->
```java
/* DO NOT ALTER OR DELETE THIS CELL! */
var result = N(100);
assert result == 131 || result == 137 || result == 139 : "A.2a.1";

```

## Aufgabe 2b

Vervollständigen Sie die Funktion `add`. Die Aufgabe von `add` ist auf Basis der `id()` einen `Student` einer Hashtabelle hinzuzufügen. Die Hashtabelle soll Verkettung verwenden, um Kollisionen handzuhaben.

- Die Hashtabelle ist als ein Feld (Array) von dynamischen Sammlungen von `Student`-Objekten implementiert.
- Die Größe der Hashtabelle soll durch die Anwendung der Funktion `N()` bestimmt werden.
- Nutzen Sie dafür die Funktion `id()`, um einen den Hashcode zu erhalten.
- Berechnen Sie die Position im Array auf Basis dieses Hashcodes (folgen Sie dazu den Hinweisen im Lehrbuch in Kapitel 15).

<!-- grade_id: cell-e3f90e02570c7b00, points: 0, grade: false, solution: true -->
```java
void add(LinkedList<Student>[] book, Student s) {
    int position = (id(s) & 0x7ffffff) % book.length;
        
    if(book[position] == null) {
        book[position] = new LinkedList<Student>();
    }
    
    book[position].addLast(s);
}

Student example = new Student("1820", "00569", "Franz Fidel Herz", "Malerei");
```

<!-- grade_id: cell-140d36b06e67d449, points: 2, grade: true, solution: false -->
```java
/* DO NOT ALTER OR DELETE THIS CELL! */
LinkedList<Student>[] test = new LinkedList[N(100)];
add(test, example);
int idx = (id(example) & 0x7ffffff) % test.length;
assert test[idx].contains(example) : "A.2b.1" ;
```

<!-- grade_id: cell-56393ba410f65ce2, points: 0, grade: false, solution: true -->
```java
// YOUR CODE

System.out.println(test.length);

if (test[(id(example)& 0x7ffffff) % test.length].contains(example)) {
   System.out.println("test"); 
}

LinkedList<Student>[] book = test;

System.out.println(book[(id(example)& 0x7ffffff) % test.length].size());
```

Output:
```
131
test
1
```

<!-- grade_id: cell-3f9aca8cf31a3cbe, points: 1, grade: true, solution: false -->
```java
/* DO NOT ALTER OR DELETE THIS CELL! */
int idx = (id(new Student("1820", "00569", "Franz Fidel Herz", "Malerei")) & 0x7ffffff) % book.length;
assert 0 < book[idx].size() && book[idx].size() <= 2 : "A.2b.2";
```

## Aufgabe 2c

Vervollständigen Sie die Funktion `get`. 
- Der return-Wert ist der gesuchte `Student`.
- Falls der gesuchte Student nicht in der übergebenen Hashtabelle zu finden ist, wird `null` zurückgegeben.

<!-- grade_id: cell-667554c40f8a0cb3, points: 0, grade: false, solution: true -->
```java
Student get(LinkedList<Student>[] book, int studentId) {
    LinkedList<Student> list = book[(studentId & 0x7ffffff) % book.length];
    for (Student s : list) {
        if (id(s) == studentId) {
            return s;
        }
    } 
    
    return null;
}

System.out.println(get(book, 182000569).name.equals("Franz Fidel Herz"));

```

Output:
```
true
```

<!-- grade_id: cell-d190c4770d1f5db9, points: 2, grade: true, solution: false -->
```java
/* DO NOT ALTER OR DELETE THIS CELL */

assert get(book, 182000569).name.equals("Franz Fidel Herz") : "A.2c.1" ;

```

## Aufgabe 3a

Vervollständigen Sie die Funktion `summary`. Die Funktion soll eine Zusammenfassung der in der Hashtabelle gespeicherten Werte, aggregiert über die Buckets und deren Belegung, liefern. Die erwartete Ausgabeform (`stdout`) ist wie folgt (als Beispiel, die konkreten Werte werden abweichen):

```
2 entries in table, 100 buckets
number of buckets with 0 entries: 0
number of buckets with 1 entries: 2
number of buckets with 2 entries: 0
number of buckets with 3 entries: 0
number of buckets with 4 entries: 0
number of buckets with 5 entries: 0
number of buckets with 6 entries: 0
number of buckets with 7 entries: 0
number of buckets with 8 entries: 0
number of buckets with 9 entries: 0
number of buckets with 10 or more entries: 0
```

<!-- grade_id: cell-68133713831798cb, points: 0, grade: false, solution: true -->
```java
void summary(LinkedList<Student>[] book) {
    int sum = 0;
    int[] amountBucket = new int[11];
    
    for (int i = 0; i < book.length; i++) {
        LinkedList<Student> list = book[i];
        int amount = 0;
        
        if (list != null) {
            sum = sum + list.size();
            amount = list.size();
        }
        
        if (amount > 10) {
            amount = 10;
        }
        amountBucket[amount]++;
    }
    
    System.out.println(sum + " entries in table, " + book.length + " buckets");
    
    System.out.println("number of buckets with 0 entries: " + amountBucket[0]);
    System.out.println("number of buckets with 1 entries: " + amountBucket[1]);
    System.out.println("number of buckets with 2 entries: " + amountBucket[2]);
    System.out.println("number of buckets with 3 entries: " + amountBucket[3]);
    System.out.println("number of buckets with 4 entries: " + amountBucket[4]);
    System.out.println("number of buckets with 5 entries: " + amountBucket[5]);
    System.out.println("number of buckets with 6 entries: " + amountBucket[6]);
    System.out.println("number of buckets with 7 entries: " + amountBucket[7]);
    System.out.println("number of buckets with 8 entries: " + amountBucket[8]);
    System.out.println("number of buckets with 9 entries: " + amountBucket[9]);
    System.out.println("number of buckets with 10 or more entries: " + amountBucket[10]);
}

```

<!-- grade_id: cell-388c652887cf8063, points: 1, grade: true, solution: false -->
```java
summary(book);
```

Output:
```
1 entries in table, 131 buckets
number of buckets with 0 entries: 130
number of buckets with 1 entries: 1
number of buckets with 2 entries: 0
number of buckets with 3 entries: 0
number of buckets with 4 entries: 0
number of buckets with 5 entries: 0
number of buckets with 6 entries: 0
number of buckets with 7 entries: 0
number of buckets with 8 entries: 0
number of buckets with 9 entries: 0
number of buckets with 10 or more entries: 0
```

## Aufgabe 3b

Erläutern Sie die Auswertung des Aufrufs von `summary(book)` in eigenen Worten in Hinblick auf folgende Aspekte einer Hashtabelle:

* Kollisionen
* Verkettung
* Entartung

<!-- grade_id: cell-9959efc06ec3cc37, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

Bei der Kollision wird 2 Elementen die gleiche Position im Array zugewiesen, wodurch man zu einer Kollisionsbehandlung gezwungen ist, die üblicherweise mit einer Verkettung oder Sondierung durchgeführt wird.
Bei unserem Beispiel ist eine Verkettung erfolgt.
Im Aufruf von der Methode summary() wird das Array durchgegangen, welches LinkedList<Student> Elemente enthält. Diese LinkedLists dienen der Verkettung um Kollisionen handeln zu können.
Wenn man nun wissen will, wieviele Elemente insgesamt in der Hashtabelle vorhanden sind, darf man das Arrary nicht einfach durchgehen, sondern man muss an jeder Position die Verkettung durchgehen, was zu einem erhöhten Zugriffsaufwand führt. Das versteht man unter Entartung, da durch die Verkettung der sonst konstante Zugriff auf das Array verloren geht.

```java

```

```java

```