# Onlineaktivität 2: Datenstrukturen (Verkettete Listen)
(Gruppenaufgabe)

Im folgenden geben wir eine einfache Implementierung einer (einfach) verketteten Liste (engl.: "Linked List") vor. 

* Eine verkettete Liste besteht aus einzelnen Listenelementen `ListItem`. Ein `ListItem` wiederum speichert einen Wert `value`, und eine Referenz `next` auf das nächste `ListItem`.
* Die `LinkedList` besitzt nur einen Zeiger auf das erste Element der Liste (`head`); auf die weiteren Elemente wird per `next` zugegriffen.

```java
class ListItem {
    
    String value;
    ListItem next;
    
    ListItem(String value){
        this.value = value;
    }
}

class LinkedList {
    ListItem head;
}

/* Create a first list and its first item */
LinkedList aList = new LinkedList();
aList.head = new ListItem("head");

/* Create another item ... */
ListItem shoulderItem = new ListItem("shoulders");
/* and append it to the list */
aList.head.next = shoulderItem;

/* Adding new items directly to the list end (via next.next and next.next.next) */
aList.head.next.next = new ListItem("knees");
aList.head.next.next.next = new ListItem("toes");

/* Each element in the list is only accessible by first referencing the one before */

System.out.println(aList.head.value);
System.out.println(aList.head.next.next.next.value);
System.out.println(aList.head.next.next.next.next);
```

Output:
```
head
toes
null
```

## Aufgabe 1

Gegeben sind die beispielhaften Listenelemente (`value, next.value`) einer `LinkedList` als Tupel von Tupeln. Jedes Listenelement ist modelliert aus dem Wert (`value`), das vom jeweiligen Listenelement gehalten wird, und dem Wert des Folgeelements (`next.value`):

```("d", null), ("a", "d"), ("b", "c"), ("h", "a"), ("f", "h"), ("c", "f")```

Beantworten Sie die folgenden Fragen:
* Welches Element ist das Erste in der Liste? Warum?
* Welches ist das Letzte? Warum?
* Ausgehend vom ersten Element `head`, was wäre die Ausgabe von `head.next.next.next.next.value`?

<!-- grade_id: cell-d5ae75562a89a0fb, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

1.) Das erste Element ist das Element mit den Attributen („b“, „c“), da es kein anderes Element mit einem next.value Wert von „b“ gibt. Es „zeigt“ also kein anderes Element auf dieses Element, warum es das Erste der Liste sein muss.

2.) Das letzte Element ist das Element mit den Attributen („d“, null), da der next.value Wert null ist. Dieses Element "zeigt" also auf kein darauffolgendes Element, warum es das Letzte der Liste sein muss.

3.) Die Ausgabe wäre: a. Das Element head.next.next.next.next hat nämlich die Attribute („a“, „d“), wobei „a“ den value-Wert darstellt, der ausgegeben wird, und „d“ den next.value Wert.


(Die gesamte LinkedList wäre: („b“, „c“) („c“, „f“) („f“, „h“) („h“, „a“) („a“, „d“) („d“, null))

## Aufgabe 2

* Gegeben sind `n` und `m` als zwei Listenelemente.
* `n` ist bereits als Listenelement in der Liste enthalten, `m` (noch) nicht.
* Erklären Sie, wie sich die folgenden Zuweisungen auf eine `LinkedList` auswirken:
  1. `n.next = n.next.next`
  2. `m.next = n.next`  
     `n.next = m`
  3. `n.next = m`  
     `m.next = n.next`

<!-- grade_id: cell-3d57536b0fb9f68d, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

1.) Das Element, das ursprünglich als n.next gespeichert war, wird mit dem Element n.next.next überschrieben. Dadurch wird das Element n.next aus der Liste entfernt und das Element n zeigt nun auf das Element, das ursprünglich als n.next.next in der Liste gespeichert war. Vereinfacht gesagt: n.next wird aus der Liste entfernt.

2.) Durch die zweite Zuweisung wird m nach n eingefügt, das durch die erste Zuweisung auf das Element verweist, das ursprünglich als n.next gespeichert war. Es wird also m zwischen n und n.next eingefügt.

3.) Durch die erste Zuweisung wird nach dem Element n das Element m in die Liste eingefügt, das durch die zweite Zuweisung auf sich selbst zeigt. Würde man diese Liste also als String ausgeben wollen, würde das zu einer Art Endlosschleife führen, da nach n.value immer wieder m.value ausgegeben werden würde.

## Aufgabe 3

* Gegeben ist die eingangs vorgestellte Implementierung von `LinkedList` und `ListItem`. 
* Erweitern Sie die Klasse `LinkedList` um die folgenden beiden Funktionen:
  * `deleteAt(pos)`: Diese Methode löscht das Element an der Position `pos` aus der Liste löschen (beginnend bei 0).
  * `reverse()`: Diese Methode gibt alle Werte der Liste als neue Liste (`LinkedList`) aus, aber in umgekehrter Reihenfolge.
  * `contains(String el)`: Diese Methode gibt `true` zurück, wenn das Stringelement `el` in der Liste enthalten ist; `false` andernfalls.


Drei Methoden für `LinkedList` sind bereits bereitgestellt:
* `append(v)`: Fügt ein neues Element ans Ende der Liste hinzu.
* `toString()`: Gibt die Werte aller Listenelemente aus. 
* `size()`: Gibt die Länge der Liste aus.

Versuchen Sie diese zu verstehen; sie können Ihnen bei der Aufgabe helfen!

<!-- grade_id: cell-268439e507cc7bec, points: 0, grade: false, solution: true -->
```java
class LinkedList {
    
    ListItem head;
        
    public void append(String v) {
        ListItem newItem = new ListItem(v);
        ListItem i = this.head;
        
        if (i == null) {
            this.head = newItem;
        } else {
            while (i.next != null) {
                i = i.next;
            }
            i.next = newItem;
        }
    }
    
    public String toString() {
        ListItem i = this.head;
        StringBuilder sb = new StringBuilder();
        if (i != null) {
            while (i.next != null){
                sb.append(i.value);
                sb.append(" ");
                i = i.next;
            }
            sb.append(i.value);
        }
        return sb.toString();
    }
    
    int size(){
        int length = 0;
        ListItem i = this.head;
        while (i != null){
            length++;
            i = i.next;
        }
        return length;
    }
    
    boolean contains(String el) {
        // YOUR CODE HERE
        
        ListItem i = this.head;
        while (!i.value.equals(el) && i.next != null) {
            i = i.next;
        }
        if (i.value.equals(el)) {
            return true;
        }
        return false;
    } 

    void deleteAt(int p){
        // YOUR CODE HERE
        
        int pos = 0;
        ListItem i = this.head;
        ListItem prev = null;
        
        if (p == 0) {
            this.head = this.head.next;
        } else {
            while (pos != p) {
                pos++;
                prev = i;
                i = i.next;
            }
            prev.next = i.next;
        }
    }
  
    LinkedList reverse() {
        // YOUR CODE HERE
        
        LinkedList revList = new LinkedList();
        int origSize = this.size();
        int revSize = 0;
        
        while (origSize > revSize) {
            ListItem i = this.head;
            for (int c = 1; c < origSize - revSize; c++) {
            i = i.next;
            }
            revList.append(i.value);
            revSize++;
        }
        return revList;
    }

}
```

Im folgenden Code wird eine Liste erstellt und dann die einzelnen Funktionen aufgerufen. Wenn Sie diese korrekt implementiert haben, sollten Sie diese Ausgabe bekommen:

```
Print all values:
head shoulders knees toes
Length:
4
List contains an element: 'knees'?
true
Delete the item at position 2:
false
Print all values in reverse order:
toes shoulders head
```

```java
LinkedList bList = new LinkedList();
bList.append("head");
bList.append("shoulders");
bList.append("knees");
bList.append("toes");

System.out.println("Print all values:");
System.out.println(bList.toString());
System.out.println("Length:");
System.out.println(bList.size());

System.out.println("List contains an element: 'knees'?");
System.out.println(bList.contains("knees"));

System.out.println("Delete the item at position 2:");
bList.deleteAt(2);
bList.toString();
System.out.println(bList.contains("knees"));

System.out.println("Print all values in reverse order:");
System.out.println(bList.reverse().toString());
```

Output:
```
Print all values:
head shoulders knees toes
Length:
4
List contains an element: 'knees'?
true
Delete the item at position 2:
false
Print all values in reverse order:
toes shoulders head
```

<!-- grade_id: cell-c1f1f84d3d11dc7b, points: 6, grade: true, solution: false -->
```java
/* DO NOT DELETE OR ALTER THIS CELL! */

assert bList.size() == 3;

```